# Python Assessment Test

### Basic python comprehension.

#### 1. What is used to define a block of code in the body of a loop in Python?
A. Parenthesis <br>
B. Curly braces <br>
C. Quotation <br>
D. Indentation <br>

<!-- grade_id: cell-05ff9bfd38f4e51d, points: 0, grade: false, solution: true -->
```python
def answer_one():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-cdf619fa136a3d9a, points: 5, grade: true, solution: false -->
```python
################################
```

#### 2. Which of the following statements is true?
A. Python is a high level programming language. <br>
B. Python is an object-oriented language. <br>
C. Python is an interpreted language. <br>
D. All of the above. <br>

<!-- grade_id: cell-c7b0914806985989, points: 0, grade: false, solution: true -->
```python
def answer_two():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-16fb09fab50fb803, points: 5, grade: true, solution: false -->
```python
################################
```

#### 3. Write a function to concatenate two strings. 

<!-- grade_id: cell-f024c61ce694f401, points: 0, grade: false, solution: true -->
```python
def answer_three(string1, string2):
    '''concatenate two strings'''
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
answer_three('hello','world')
```

<!-- grade_id: cell-f74436aeb8f4b81e, points: 5, grade: true, solution: false -->
```python
"""Check that answer_three returns the correct output for several inputs"""
assert answer_three('', 'hi') == 'hi'
assert answer_three('hi','there') == 'hithere'
```

#### 4. What is the output of the following code? `print(3 >= 3)`
A. True <br>
B. 3>=3 <br>
C. False <br>
D. None<br>

<!-- grade_id: cell-f867a7d544cb7831, points: 0, grade: false, solution: true -->
```python
def answer_four():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2fd19eba3cfb4730, points: 5, grade: true, solution: false -->
```python
################################
```

#### 5. if `a='1'` and `b=2` then `a*b` would return?
A. 2 <br>
B. 1 <br>
C. 11 <br>
D. Error <br>

<!-- grade_id: cell-af9400f185ab222c, points: 0, grade: false, solution: true -->
```python
def answer_five():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9ed8679aaf8b2b7c, points: 5, grade: true, solution: false -->
```python
################################
```

#### 6. What is the output of the following program? `print((1, 2) + (3, 4))`
A. (1, 2, 3, 4) <br>
B. (4, 6)<br>
C. ((1, 2), (3, 4))<br>
D. error<br>

<!-- grade_id: cell-e93851f3589e3035, points: 0, grade: false, solution: true -->
```python
def answer_six():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d95aa2d83df11d7e, points: 5, grade: true, solution: false -->
```python
################################
```

#### 7. Write a function to print the last four elements of a list.

<!-- grade_id: cell-d7392ef89a5e0dd8, points: 0, grade: false, solution: true -->
```python
def answer_seven(my_list):
    '''print the last four elements'''
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

your function should return `['t','h','o','n']` for `answer_seven(['P','y',t','h','o','n'])` Check that it does:

```python
answer_seven(['P','y','t','h','o','n'])
```

<!-- grade_id: cell-3897fd7d5954f8a7, points: 5, grade: true, solution: false -->
```python
"""Check that answer_seven returns the correct output for several inputs"""
assert answer_seven(['w','o','r','l','d']) == ['o','r','l','d']
assert answer_seven(['h','e','l','l','o','o','o']) == ['l','o','o','o']
```

#### 8. Suppose a tuple named test contains 5 elements. How can you set the 3rd element of the tuple to `Python`?
A. `test(2) = 'Python'` <br>
B. `test[2] = 'Python'` <br>
C. `test[3] = 'Python'` <br>
D. `Elements of tuple cannot be changed` <br>

<!-- grade_id: cell-11a9ce99f17c0b43, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b06655eddad1371d, points: 5, grade: true, solution: false -->
```python
################################
```

#### 9. If `n='9'`, what is the value type of n? 
A. integer <br>
B. string <br>
C. tuple <br>
D. list <br>

<!-- grade_id: cell-0ed7a1b1dc77c5b6, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-46b764a9e1376dd5, points: 5, grade: true, solution: false -->
```python
################################
```

#### 10. Which function is used to open a file for reading in python?
A. `fopen(filename,'r')` <br>
B. `open(filename,'r')`<br>
C. `openfile(filename,'r')`<br>
D. `open_file(filename,'r')`<br>

<!-- grade_id: cell-5c68c0e1850b9b5b, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-44b97cc25f62f02b, points: 5, grade: true, solution: false -->
```python
################################
```

#### 11. What keyword is used to create a function?
A. `funtion` <br>
B. `func` <br>
C. `def` <br>
D. `import` <br>

<!-- grade_id: cell-6ac096ed5b8c73b0, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ece82f567fbb6b4a, points: 5, grade: true, solution: false -->
```python
################################
```

#### 12. which command returns the first element of the list `dat`?
A. `dat[0]` <br>
B. `dat[1]` <br>
C. `dat[-1]` <br>
D. `dat[1:]` <br>

<!-- grade_id: cell-6d4936a19e1940a9, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-764646aa7bee0b2b, points: 5, grade: true, solution: false -->
```python
################################
```

#### 13. In python, what is the correct method to load the math module?
A. `include math` <br>
B. `import math` <br>
C. `using math` <br>
D. `with math` <br>

<!-- grade_id: cell-ec18f87055d44c7b, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3d09beedf3a90706, points: 5, grade: true, solution: false -->
```python
################################
```

#### 14. if `var1='Python'` what does `var1[-4:]` return?
A. 'Py' <br>
B. 'Python' <br>
C. 'hon' <br>
D. 'thon' <br>

<!-- grade_id: cell-e58cada76ac96801, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6cc36b400f239228, points: 5, grade: true, solution: false -->
```python
################################
```

#### 15. Which of the following statement is true?
A. A set is an unordered collection of items. <br>
B. Elements of set are unique. <br>
C. You can change elements of set unlike tuple. <br> 
D. All of the above. <br>

<!-- grade_id: cell-078605af72e43d98, points: 0, grade: false, solution: true -->
```python
def answer_fifteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1cf00a488f21e988, points: 5, grade: true, solution: false -->
```python
################################
```

#### 16. What command will open a database connection in pandas?
A. `pd.read_sql` <br>
B. `pd.read_odbc` <br>
C. `pd.read_database` <br>
D. `pd.open_database` <br>

<!-- grade_id: cell-8f3ad757cccacdfb, points: 0, grade: false, solution: true -->
```python
def answer_sixteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7703fed1e39aaf30, points: 5, grade: true, solution: false -->
```python
################################
```

#### 17. what command will return a unique group of a list (list1)?
A. `distinct(list1)` <br>
B. `set(list1)` <br>
C. `unique(list1)` <br>
D. `append(list1)` <br>

<!-- grade_id: cell-b79863cf86330a21, points: 0, grade: false, solution: true -->
```python
def answer_seventeen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-87598c5d5495175f, points: 5, grade: true, solution: false -->
```python
################################
```

#### 18. How do you slice `['b', 'c']` from the list `t = ['a', b', 'c', 'd', 'e]`?
A. `t[1:3]` <br>
B. `t[1:]` <br>
C. `t[1,2]` <br>
D. `t[2]` <br>

<!-- grade_id: cell-79393eda8ab00220, points: 0, grade: false, solution: true -->
```python
def answer_eighteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9c4702b98f03ab25, points: 5, grade: true, solution: false -->
```python
################################
```

#### 19. How can you change `num = {‘one’: 1, ‘two’: 3}` to `num = {‘one’: 1, ‘two’: 2}`? 
A. `num[2] = 'two'` <br>
B. `num[1] = 'two'` <br>
C. `num['two'] = 2` <br>
D. `num['two'] = '2'` <br>

<!-- grade_id: cell-188cddfb01e1c1b7, points: 0, grade: false, solution: true -->
```python
def answer_nineteen():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2c6962bd26cb1460, points: 5, grade: true, solution: false -->
```python
################################
```

#### 20. What is used to add a new element at the end of a list?
A. `append()` <br>
B. `add()` <br>
C. `extend()` <br>
D. `pop()` <br>

<!-- grade_id: cell-1e9ce5399e6caa91, points: 0, grade: false, solution: true -->
```python
def answer_twenty():
    '''return the letter of the correct answer'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e297fe00da2a6b6c, points: 0, grade: true, solution: false -->
```python
################################
```

```python

```